<!-- grade_id: cell-9f1860a33162d9ba, points: 0, grade: false, solution: true -->
```python
def squares(n):
    return [i**2 for i in range(1,n+1)]
```

<!-- grade_id: cell-6ac8a463ab2d2354, points: 2, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(2) == [1,4]
```

```python

```