<!-- grade_id: cell-f016563141f61d0c, points: 0, grade: false, solution: true -->
```python
def squares(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-fd15be75ca0aff93, points: 2, grade: true, solution: false -->
```python
assert squares(1) == [1]
assert squares(2) == [1,4]
```

```python

```